---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: f3bb40b3db
xp: 100
```

This is the first time to use datacamp and test code function

`@instructions`
This is the first time to use datacamp and test code function

`@hint`
This is the first time to use datacamp and test code function

`@pre_exercise_code`
```{python}
import torch as th
```

`@sample_code`
```{python}
N/A
```

`@solution`
```{python}
N/A
```

`@sct`
```{python}

```
